---
title: DELETE
description: Insert the chapter description here
---

## DELETE

```yaml
type: NormalExercise
key: 16f6ff0c7d
lang: python
xp: 100
skills:
  - 2
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
